---
title: "The normal distribution"
output: 
  learnr::tutorial:
    progressive: true
    allow_skip: true
runtime: shiny_prerendered
---

```{r setup, include=FALSE,message=FALSE}
library(learnr)
library(learnrhash)
library(tidyverse)
library(gradethis)
library(openintro)
library(infer)
tutorial_options(
  # use gradethis for checking
  exercise.checker = gradethis::grade_learnr
  )
knitr::opts_chunk$set(echo = FALSE)

mcdonalds <- fastfood %>%
  filter(restaurant == "Mcdonalds")
dairy_queen <- fastfood %>%
  filter(restaurant == "Dairy Queen")
dqmean <- mean(dairy_queen$cal_fat)
dqsd   <- sd(dairy_queen$cal_fat)
```

## Logistics

This lab will occur remotely and in groups of three. You will work with your final project team so that if you finish early you can work on your project together. For those of you participating synchronously, you will find the Zoom room information on [Moodle](https://moodle.smith.edu/). I will assign you to a breakout room with your partner(s) from last week. 

For those participating syncronously I will be available in the main Zoom room to answer questions. If you have a question or technical problem, click the "Ask for Help" button (it looks like a question mark) in the meeting controls, and I will be alerted to join your breakout room.  

For those of you participating asynchronously, alert me to technical challengs over Slack DM, and I will get back to you as soon as possible. For questions about the content of the lab, please sign up for an office hour time block. 

Each of you should be writing and running code, examining output, and answering the exercises throughout the lab. However, you only need to turn in one final lab report. To be clear, everyone submits files to Moodle, but the files can be the same within a group. Today one of you should be the main recorder of answers in the lab document. You will share this document with your teammates. As you work it may be helpful to share your screen. It may be helpful to share your screen. 

You all should also feel free to ask and answer questions amongst yourselves via Zoom if participating synchronously or via Slack if participating asynchronously. Please note at the end of the lab document who you consulted for help.

## Getting Started

In this lab, you'll investigate the probability distribution that is most central
to statistics: the normal distribution.  If you are confident that your data are 
nearly normal, that opens the door to many powerful statistical methods.  Here 
we'll use the graphical tools of R to assess the normality of our data and also 
learn how to generate random numbers from a normal distribution.

### Load packages

In this lab, we will explore and visualize the data using the **tidyverse** suite of packages as well as the **openintro** package.

Let's load the packages.

```{r load-packages, exercise = T}
library(tidyverse)
library(openintro)
```


### The data

This week you'll be working with fast food data.  This data set contains data on
515 menu items from some of the most popular fast food restaurants worldwide.
Let's take a quick peek at the first few rows of the data.

Either you can use `glimpse` like before, or `head` to do this.

```{r load-data, exercise = T}
glimpse(fastfood)
head(fastfood)
```

You'll see that for every observation there are 17 measurements, many of which are
nutritional facts.

You'll be focusing on just three columns to get started: restaurant, calories, 
calories from fat.

Let's first focus on just products from McDonalds and Dairy Queen.

```{r male-female, exercise = T}
mcdonalds <- fastfood %>%
  filter(restaurant == "Mcdonalds")
dairy_queen <- fastfood %>%
  filter(restaurant == "Dairy Queen")
```

Make plots to visualize the distributions of the amount of calories
from fat of the options from these two restaurants.  How do their centers, shapes,
and spreads compare?

```{r dghqyduwvsgyzshm, exercise = TRUE}
ggplot(mcdonalds, aes(x = ___)) + ___
ggplot(dairy_queen, aes(x = ___)) + ___
```

```{r dghqyduwvsgyzshm-solution}
ggplot(mcdonalds, aes(x = cal_fat)) + geom_histogram()
ggplot(dairy_queen, aes(x = cal_fat)) + geom_histogram()
```

```{r dghqyduwvsgyzshm-check}
# check code
gradethis::grade_code()
```

## The normal distribution

In your description of the distributions, did you use words like *bell-shaped* 
or *normal*?  It's tempting to say so when faced with a unimodal symmetric 
distribution.

To see how accurate that description is, you can plot a normal distribution curve 
on top of a histogram to see how closely the data follow a normal distribution. 
This normal curve should have the same mean and standard deviation as the data. 
You'll be focusing on calories from fat from Dairy Queen products, so let's store
them as a separate object and then calculate some statistics that will be referenced 
later. 

```{r female-hgt-mean-sd, exercise = T}
dqmean <- mean(dairy_queen$cal_fat)
dqsd   <- sd(dairy_queen$cal_fat)
```

Next, you make a density histogram to use as the backdrop and use the `lines` 
function to overlay a normal probability curve. The difference between a 
frequency histogram and a density histogram is that while in a frequency 
histogram the *heights* of the bars add up to the total number of observations, 
in a density histogram the *areas* of the bars add up to 1. The area of each bar
can be calculated as simply the height *times* the width of the bar. Using a 
density histogram allows us to properly overlay a normal distribution curve over
the histogram since the curve is a normal probability density function that also
has area under the curve of 1. Frequency and density histograms both display the
same exact shape; they only differ in their y-axis. You can verify this by
comparing the frequency histogram you likely constructed earlier and the density
histogram created by the commands below.

```{r hist-height, exercise = T}
ggplot(data = dairy_queen, aes(x = cal_fat)) +
        geom_blank() +
        geom_histogram(aes(y = ..density..)) +
        stat_function(fun = dnorm, args = c(mean = dqmean, sd = dqsd), col = "tomato")
```

After initializing a blank plot with `geom_blank()`, the `ggplot2` package (within
the `tidyverse`) allows us to add additional layers. The first layer is a density 
histogram. `..density..` is a special term that tells `ggplot` to calculate the density for the specified variable (here `cal_fat`) and use it to define the value on chosen axis. The second layer is a statistical function -- the density of the normal
curve, `dnorm`. We specify that we want the curve to have the same mean and 
standard deviation as the calories from fat in the dairy queen data. The argument `col` simply sets
the color for the line to be drawn. If we left it out, the line would be drawn in
black.


## Evaluating the normal distribution

Eyeballing the shape of the histogram is one way to determine if the data appear
to be nearly normally distributed, but it can be frustrating to decide just how 
close the histogram is to the curve. An alternative approach involves 
constructing a normal probability plot, also called a normal Q-Q plot for 
"quantile-quantile".

```{r qq, exercise = T}
ggplot(data = dairy_queen, aes(sample = cal_fat)) + 
  geom_qq() + geom_qq_line(alpha = 0.7, color = "red", linetype = "dashed")
```

This time, you can use the `geom_qq()` layer, while specifying that you will 
be creating a Q-Q plot with the `stat` argument. It's important to note that here,
instead of using `x` instead `aes()`, you need to use `sample`.

The x-axis values correspond to the quantiles of a theoretically normal curve 
with mean 0 and standard deviation 1 (i.e., the standard normal distribution).  The 
y-axis values correspond to the quantiles of the original unstandardized sample 
data. However, even if we were to standardize the sample data values, the Q-Q
plot would look identical. A data set that is nearly normal will result in a
probability plot where the points closely follow a diagonal line.  Any deviations
from normality leads to deviations of these points from that line.  `geom_qq_line()` provides a guide line to help you assess whether points deviate from a line.

The plot for Dairy Queen's calories from fat shows points that tend to follow the
line but with some errant points towards the upper tail.  You're left with the same
problem that we encountered with the histogram above: how close is close enough?

A useful way to address this question is to rephrase it as: what do probability 
plots look like for data that I *know* came from a normal distribution?  We can 
answer this by simulating data from a normal distribution using `rnorm`.

```{r sim-norm, exercise = T}
sim_norm <- data.frame(sim_val = rnorm(n = nrow(dairy_queen), mean = dqmean, sd = dqsd))
```

The first argument indicates how many numbers you'd like to generate, which we 
specify to be the same number of menu items in the `dairy_queen` data set using the 
`nrow()` function.  The last two arguments determine the mean and standard 
deviation of the normal distribution from which the simulated sample will be 
generated.  You can take a look at the shape of our simulated data set, `sim_norm`, 
as well as its normal probability plot.

2.  Make a normal probability plot (qq-plot) of `sim_norm`.  Do all of the points fall on 
    the line?  How does this plot compare to the probability plot for the real 
    data? (Since `sim_norm` is not a dataframe, it can be put directly into the 
    `sample` argument and the `data` argument can be dropped.)
    
```{r vuqymqrklhvinkjz-setup}
sim_norm <- data.frame(sim_val = rnorm(n = nrow(dairy_queen), mean = dqmean, sd = dqsd))
```
    
```{r vuqymqrklhvinkjz, exercise = TRUE}
ggplot(data = ___, aes(sample = ___)) + 
  geom_qq() + geom_qq_line(alpha = 0.7, color = "red", linetype = "dashed")
```


```{r vuqymqrklhvinkjz-solution}
ggplot(data = sim_norm, aes(sample = sim_val)) + 
  geom_qq() + geom_qq_line(alpha = 0.7, color = "red", linetype = "dashed")
```

```{r vuqymqrklhvinkjz-check}
# check code
gradethis::grade_code()
```

Even better than comparing the original plot to a single plot generated from a 
normal distribution is to compare it to many more plots using the following 
function. It shows the Q-Q plot corresponding to the original data in the top
left corner, and the Q-Q plots of 8 different simulated normal data.

```{r qqnormsim, exercise = T}
qqnormsim(sample = cal_fat, data = dairy_queen)
```

Using the same technique, determine whether or not the calories from McDonald's menu appear to come from a normal distribution.

```{r gbmknvwczbalsxme, exercise = TRUE}
qqnormsim(sample = ___, data = ___)
```

```{r gbmknvwczbalsxme-solution}
qqnormsim(sample = cal_fat, data = mcdonalds)
```

```{r gbmknvwczbalsxme-check}
# check code
gradethis::grade_code()
```


## Normal probabilities

Okay, so now you have a slew of tools to judge whether or not a variable is 
normally distributed.  Why should you care?

It turns out that statisticians know a lot about the normal distribution.  Once 
you decide that a random variable is approximately normal, you can answer all 
sorts of questions about that variable related to probability.  Take, for 
example, the question of, "What is the probability that a randomly chosen Dairy 
Queen product has more than 600 calories from fat?"

If we assume that the calories from fat from Dairy Queen's menu are normally 
distributed (a very close approximation is also okay), we can find this probability
by calculating a Z score and consulting a Z table (also called a normal probability
table).  In R, this is done in one step with the function `pnorm()`.

```{r pnorm, exercise = T}
1 - pnorm(q = 600, mean = dqmean, sd = dqsd)
```

Note that the function `pnorm()` gives the area under the normal curve below a 
given value, `q`, with a given mean and standard deviation.  Since we're 
interested in the probability that a Dairy Queen item has more than 600 calories
from fat, we have to take one minus that probability.

Assuming a normal distribution has allowed us to calculate a theoretical 
probability.  If we want to calculate the probability empirically, we simply 
need to determine how many observations fall above 600 then divide this number 
by the total sample size. This is what we did when we were computing p-values ourselves.

```{r probability, exercise = T}
dairy_queen %>% 
  filter(cal_fat > 600) %>%
  summarise(percent = n() / nrow(dairy_queen))
```

Although the probabilities are not exactly the same, they are reasonably close. 
The closer that your distribution is to being normal, the more accurate the 
theoretical probabilities will be.

## Submit checkpoint

```{r context="server"}
learnrhash::encoder_logic(strip_output = T)
```

```{r encode, echo=FALSE}
learnrhash::encoder_ui(
  ui_before = div(strong("Submit your hash in the form below."), br(), br()),
  ui_after  = learnrhash::iframe_ui(
    src = "https://docs.google.com/forms/d/e/1FAIpQLSda78rAQWziQ-L6eDanGoDAhcfoi6JU_dg-c-6A9rrEXdvJwg/viewform", ## change link, include name
    width="900px", height= "1000px"
  )
)
```

## Creating a reproducible lab report

For the rest of this lab you will R Markdown to create a reproducible lab report. 
In RStudio, go to New File -> R Markdown... Then, choose From Template and then choose `Lab Report` from the list of templates. Make sure to name the document appropriately and pick a location for the file where you know how to find it on your computer.

See the following video describing how to get started with creating these 
reports for this lab, and all future labs:

[**Basic R Markdown with an OpenIntro Lab**](https://www.youtube.com/watch?v=Pdc368lS2hk)
<iframe width="560" height="315" src="https://www.youtube.com/embed/Pdc368lS2hk" frameborder="0" allowfullscreen></iframe>

**Note: This video was made using `oilabs`. We will NOT be using `oilabs`. Be sure to pick the Lab Report template that goes with `openintro` as shown in screenshot above. Make sure you have `library(openintro)` in your first chunk not `library(oilabs)`.**

## Questions for Lab Report

Answer in an Rmd file based on the lab report template. Remember you will need to load appropriate packages, set a seed, and load data. 

1. It would be nice if multiple simulated datasets were all on the same plot to better judge whether our data fits within the simulated data. We repurpose some of the code that makes up the `qqnormsim` function. To see the source code for a function (what's going on behind the scenes when you use a function), you just specify the name without any parentheses or arguments.

```{r, eval = F}
qqnormsim
```

We'll grab the parts we need, and adjust the parts we want to change. It is common practice to build off of pre-existing code when we develop new functions. It is not crucial that you understand every line of this chunk; just focus on the resulting plot. 

```{r, eval = F}
y <- dairy_queen$cal_fat
simnorm <- rnorm(
  n = length(y) * 8, mean = mean(y),
  sd = sd(y)
)
df <- data.frame(x = c(y, simnorm), plotnum = rep(c(
  "data",
  "sim 1", "sim 2", "sim 3", "sim 4", "sim 5", "sim 6",
  "sim 7", "sim 8"
), each = length(y)))

original <- df %>% filter(plotnum == "data")
df <- df %>% filter(plotnum != "data")

ggplot(data = df, aes(sample = x, group = plotnum)) + geom_qq(alpha = 0.6) + geom_qq(data = original, aes(sample = x), col = "red", size = 2)
```

2. In the plot above does the normal probability plot for the calories from fat in Dairy Queen items look similar to the plots created for the simulated data?  That is, do the plots provide evidence that the calories from fat are nearly normal? Explain.

3. Write out a probability question that you would like to answer about any 
    of the restaurants in this dataset that would take advantage of `pnorm`.  Calculate that probability.

4.  Write out a question that you would like to answer about any 
    of the restaurants in this dataset that would take advantage of `qnorm`.  Calculate that quantile.

## Deliverables

Make sure you have submitted your hash to the Google Form for the first part of the lab.

When you are finished editing your Markdown document click the "Knit" button and choose "Knit to HTML" in the top left corner of RStudio. This will run all of your code and create a formatted document of the output. If you get an error, it means something in your Markdown file isn't right, either an error in code or some error in formatting. Call me into your breakout room, and we will troubleshoot.

Submit your Markdown document and knitted file to [Moodle](https://moodle.smith.edu/) as:

LastName-LastName-LastName-L-10.Rmd  (add a third last name if applicable)

LastName-LastName-LastName-L-10.html

*Due*: Monday (beginning of class time, Eastern time)

* * *

<a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/"><img alt="Creative Commons License" style="border-width:0" src="https://i.creativecommons.org/l/by-sa/4.0/88x31.png" /></a><br />This work is licensed under a <a rel="license" href="http://creativecommons.org/licenses/by-sa/4.0/">Creative Commons Attribution-ShareAlike 4.0 International License</a> and was adapted and learnr-ified by Sara Stoudt.